# Lab 1. Introduction to Python

#### Table of contents

1. Overview
2. Jupyter notebook
3. A quick introduction to `Python`
 - 3.1. Numbers
 - 3.2. Operations
 - 3.3. Variables
 - 3.4. Datatype
 - 3.5. Fuctions I
 - 3.6. Conditions
 - 3.7. Operators
 - 3.8. `while` loop
 - 3.9. `for` loop
 - 3.10. Functions II
 - 3.11. Lists
 
4. Quiz 1

## 1. Overview

During this Lab session we will learn the programming language __Python__.  The goal is to become familiar with the fundamental commands. Most of this lab is reading but don't forget to complete the quiz in section 4 and submit your assignment before the deadline.

## 2. Jupyter notebook

__Jupyter Notebook__ is an open-source web application that allows you to create and share documents that contain live code, equations, visualizations and narrative text. Uses include: data cleaning and transformation, numerical simulation, statistical modeling, data visualization, machine learning, and much more.

This is a notebook, double click on the text and you will discover the `code` behind.
To get back to the formated text type `shift+enter`.
This section is locked as _read only_ so you cannot modify its content but later you will encounter other sections that are enabled for you to write code in it.
A notebook is made of __cells__ and each cell can have one of the 3 types listed below: 

- code (can be executed as Python code)
- markdown (the one used for this cell, i.e. formatted text)
- raw text

The type of each cell is displayed on the dropdown menu here showing `Markdown` on the menu bar above the notebook window (if you select this cell).
To submit any selected cell to the kernel to interpret it (as code, markdown or raw) just type `shift+enter`. 
__To prevent any error in this notebook, a good practice is to interpret a cell once you have read it.__
At the end of the reading you should have interpreted all cells.

You can refer to the documentation of Jupyter to learn more about notebooks and we will discover further advanced options throughout the laboratory sessions of the class.

## 3. A quick introduction to `Python`

Not everything is interpreted by Python. For example let's try to interpret the sentence written below in a code cell (remember that to interpret the cell, select it and type `shift+enter`).

```python
Hi! ny name is Nicolas.
```

Python is returning and error of invalid syntax. We will learn in the following what Python can interpret i.e. the __Python syntax__.

### 3.1. Numbers

Numbers in Python are interpreted directly. Try to run the number 1 below. Just select the cell and type `shift+enter` to run the kernel.

```python
1
```

You see that the kernel returns the number 1, no error is reported.

Note that a Jupyter notebook represents an _interactive_ interpreter of Python code meaning that it often returns the result of the operation you input. This won't be always the case if you use Python in the more traditional manner by writing your code in a file (more on this later).

### 3.2. Operations

Addition, substraction, multiplication and division can be performed with the operators +, -, / and *. Here are some examples, try to guess the result then run the kernel to verify the answer.

```python
3+4
```

```python
-2+93
```

```python
9.5+2
```

```python
2/3+1
```

The result of a division can be separated between its integer and residual part using the operators `//` and `%` (called the modulo operator).

```python
10/3
```

```python
10//3
```

```python
10%3
```

Now try your own operations in the cell below which is enabled for you. Don't forget to hit `shift+enter` to interpret the result.

```python

```

### 3.3. Variables

Variables are central to all programming languages. A variable is a link to a data store in the computer that you will be able to use and access multiple times in order to do computations. You can imagine your computer as a giant shelf with many drawers and each drawer contains data, some contain variables.

In Python to define a variable you just need to write `name_of_the_variable = value`. Variable must fullfil basic rules:

- Name can only contain alphanumeric characters and underscore symbols `_`
- The name of the variable cannot start with a number
- Python is case sensitive meaning that `size` and `SIZE` correspond to two different variables

Beyond these strict rules, there are some good programming habits such as writing variables in lowcase and if they require more than one word, filling the space with the underscore symbol. It is always a good practice to name variables explicitely such that another coder should understand your code. Another alternative is to join words without space starting subsequent words with a capital. Below are some examples of _properly defined_ variables.

```python
my_size = 173
```

```python
mySize = 173
```

```python
my_age = 36
```

```python
myAge = 36
```

If you execute the variable definition above, the Python shell does not return anything meaning there is no error and the variable was well stored in the memory of the computer. We say that we have __assigned__ the value to the variable. Now if you want to show the value of a previously defined variable, just write its name in the interactive cell and Python will return its __value__. We note that this won't return anything if the code is executed in the more common non-interactive file.

```python
my_size
```

We can now perform operations with the variables as we have learnt before with raw numbers. For example, let's try to double the value of the variable `my_size`.

```python
my_size = my_size * 2
```

```python
my_size
```

The value of the variable was well re-assigned to twice its original value. We can also define a new variable with the value of twice my size.

```python
my_size2 = mySize * 2
```

```python
my_size2
```

Some keywords are already reserved by Python for some special usage. Here is the list for Python 3 (the version we are using now).

<img src="python3_commands.png" alt="Python 3 pre-defined commands" width=400/>

You can play with variables in the code cell below:

```python

```

### 3.4. Datatype

So far we have been working with numbers only. Python needs to know what datatype he is dealing with in order to perform operations. We will review below some of the datatype Python can handle.

#### Interger numbers

This type is called `int` in Python. For example:

```python
3
```

#### Float numbers

Float numbers have decimals and are called `float` in Python.

```python
3.52
```

#### Strings

One can define a string (`str`) in three ways:

- between double quote symbols `"This is a string"`
- between single quote symbols `'This is a string'`
- between triple double quote symbols `"""This is a string"""` (less used)

We can store a string in a variable as for example:

```python
name = 'Nicolas'
```

```python
name
```

If your string contains quote symbols, you might get the follwoing error:

```python
whose_name = 'Nicolas's'
```

However, you can escape the quote with a backslash `\`.

```python
whose_name = 'Nicolas\'s'
```

```python
whose_name
```

```python
string = '\"Silence is gold\"'
```

```python
string
```

Triple double quotes are less common and can be used to avoid having to escape quotes and double quotes and to enable multiline string.

```python
long_variable = """This is an
example of long text including "quotes"
and multiple lines"""
```

```python
long_variable
```

Note that the line return was replaced by the `\n` symbol. This symbol means line return in many programmong languages. We will get back to this later in the tutorials. The cell below is for you to test some of what you just learnt.

```python

```

### 3.5. Functions I

Functions execute an ensemble of instructions. If you need to execute the same ensemble of instruction several times you just call the function by its name and avoid repeating multiple times the same lines of code.
Most of the functions require parameters to work on it that you can provide as an input.
Functions are used with the following syntax: `name_of_the_function(parameter_1, parameter_2,...,parameter_n)`

- You start with the name of the function
- Parameters are surounded by parenthesis. If no parameter is required you still need to write the parenthesis symbols i.e. `name_of_the_function()`

Below we present some basic useful functions. We will learn how to create our own functions later.

#### The function `type()`

In the introduction above we discussed datatype. The power of Python is that it understands automatically what is the type of a variable you asign. But it can be useful to know what is the type of a given variable.

```python
type(my_age)
```

```python
type(3.4)
```

```python
type(name)
```

#### The function `print()`

This function will write to the screen whatever you pass to it. Jupyter notebooks are interactive programs and you will find the need at one point to write programs outside notebooks. In non-interactive mode, writing down the name of a variable won't return anything. Therefore you will use the `print()` function to return the value of a variable. Moreover, remember that the interactive shell returned the end of line character `\n` instead of actually going to the next line. This won't be the case with `print()`.

```python
print(name)
```

```python
print(long_variable)
```

```python
print("My name is", name)
```

In the later example, I passed 2 parameters to the function: 

- a string `"My name is"`
- the variable `name`

separated by a comma. The function print returns the value of these variables separated by a space. Print can take more than 2 parameters.

```python
print("My name is", name, "I am", mySize, "cm height.")
```

### 3.6. Conditions

Condition will enable the execution of one or several instructions in a case or one or several other in another case. This represents the building blocks of traditional programming.

#### Minimal condition `if`

Let say we want the program to perform an action if a variable `a` is positive and another action if `a` is negative. We should write something like (the cell below can be modified for you to try different values of `a`):

```python
# Example of a condition
a = 5
if a > 0:
    print('The variable \'a\' is positive')
```

Here is the explanation of the above code line-by-line:

1. the first line is a comment starting with the character `#`, the line is ignored
2. we assign the value 5 to the variable `a`
3. here is the condition including:
 - the keyword `if` (that you can find in the variable list dedicated to Python presented above)
 - the condition `a > 0`
 - the colon symbol `:` to terminate the condition statment
4. the command to execute if the condition is fulfilled. This line constitutes the instruction block and can be multiple lines.

It is critical to note that the instruction block is __indented__ by several spaces (here 4 spaces which represents a good practice). This indentation defines the instruction block and enables a clear lecture of the code. In various other programming languages the instruction block is surrounded by brackets (as well as the condition). For example in c++ we would write `if (a > 0) {printf(...)};`

#### The `else` condition

We might want the code to execute another command if the condition `if` is not fullfiled. This is where the `else` condition comes in. 

```python
age = 21
if age >= 18:
    print('You are an adult')
else:
    print('You are a minor')
```

Note that Python executes either one or the other condition but not both. Moreover the command `else` must be at the same indentation level as the command `if` and it ends with the column symbol because it is a condition. Try other values of the valiable `age`.

#### The `elif` condition

This command is a contraction of `else if`.

```python
a = -1
if a > 0:
    print('a is positive')
elif a < 0:
    print('a is negative')
else:
    print('a is zero')
```

Similarly to `else`, `elif` is at the same level of indentation as if. However between `elif` and the colon symbol, there is a new condition. Python executes the above code as follow:

- First Python checks if `a` is positive, if yes, the code enters the condition block, prints the result and stops
- If no, Python checks if `a` is negative, if yes, the code enters the second block, prints the result and stops
- If no, Python enters the last block end prints the results that `a` is zero

You can add as many `elif` as you want after a `if` however, there can only be one `else` condition.
Note that it is possible to create nested conditions and you can try it out in the next code cell.

```python

```

### 3.7. Operators

Conditions introduce new operators for comparison:

- `<` striclty less than
- `>` strictly greater than
- `<=` less or equal
- `>=` greater than or equal
- `==` equal to
- `!=` different than

Note that the comparative operator equal to `==` is different than the assignment operator `=`. Using one instead of the other is a very common programming mistake. 

#### Booleans

Here are some examples:

```python
a = 0 # Here we asign the value of 0 to variable a
a == 5 # Here we compare the value of variable a to 5
```

```python
a > -8
```

```python
a != 1
```

The shell either returns `True` or `False`. These are the two possible values of a new datatype called boolean (`bool`). Note that these two keyword start with a capital letter.

```python
type(a != 1)
```

Here is a minimal code to modify a boolean value.

```python
age = 18
adult = False
print("Before condition:",adult)
if age >= 18:
    adult = True
print("After condition:",adult)
```

After running the above few lines of code, the variable `adult` becomes `True`.

#### Keywords `and`, `or` and `not`

Often we need to test for more than one condition. For example if we want to check if a variable is within a range of values. Here is a sample code.

```python
a = 5
if a >= 2:
    if a <= 8:
        print("a is within the range.")
    else:
        print("a is not within the range.")
else:
    print("a is not within the range.")
```

This works however it is a bit heavy and one can simplified this short code using the keyword `and`.

```python
if a >= 2 and a <= 8:
    print("a is within the range.")
else:
    print("a is not within the range.")
```

Now we can check if `a` is not within the range as:

```python
if a < 2 or a > 8:
    print("a is not within the range.")
else:
    print("a is within the range.")
```

Finally, the `not` keyword can be used to check the opposite of a condition. For example, the condition `not a == 0` is equivalent to `a != 0`.

### 3.8. `while` loop

Loops can be used to repeat some instructions in a code. Programers are lazy and usually try to minimize the number of line of code to write. For example, you might want to print to the screen the multiplication table of the number 7 from 7 to 70. With our present knowledge you would have to write the (long) code below:

```python
print('1*7 =',1*7)
print('2*7 =',2*7)
print('3*7 =',3*7)
print('4*7 =',4*7)
print('5*7 =',5*7)
print('6*7 =',6*7)
print('7*7 =',7*7)
print('8*7 =',8*7)
print('9*7 =',9*7)
print('10*7 =',10*7)
```

Try to write the equivalent code using a variable `nb` instead of the _hardcoded_ number 7. This would make it easy to print another multiplication table by changing only one line of code (instead of all in the example above).

```python
# Write your code here
```

The `while` loop is common to most of the programming language and is used to repeat a block of instruction until the while condition is fulfilled. In the following example, we will create a variable that will be incremented inside the instruction block i.e. a counter. While the counter will be less than 10, the instruction block will be executed. Here we go (explanation in the following block).

```python
nb = 7
i = 0

while i < 10:
    print(i + 1,'*',nb,'=',(i+1)*nb)
    i = i + 1
```

This is much shorter code than writing 10 times a (similar) line of code and can be easily extended to any multiplication table and up to any (reasonable) factor. Remember that in programming if you are writing several time a similar line of code, there is probably a way to shorten your code.

And now the line-by-line explanation:

1. We asign the number 7 to the variable `nb`
2. We asign the number 0 to the variable `i`. Here we initialize the counter commonly called `i`.
3. Just an empty line to make the code more clear
4. The condition __while i is strictly less than 10__. Do not forget the colon symbol at the end of the line to terminate the condition.
5. The print function here contains 5 arguments. Note that in programming we often start loops with the zero index therefore if we first want to print the multiplication 7$\times$1 we must add 1 to the counter `i`. The second argument is a string containing a star symbol, the third the variable `nb`, the fourth the string `=` and the last argument is the result of the multiplication.
6. Finally, we increment the counter `i`. During the first execution of the block, `i` is equal to 0 and is incremented to 1. 

The condition `while` is fulfilled and the block is executed over and over again until the variable `i` is equal to 9. After printing `10 * 7 = 70`, the increment brings `i` to 10 and the condition `while` is no longer true and the loop exits.

We note that it is common to simplify the increment `i = i+1` with `i += 1` (told you programers were lazy). Try it out in the previous code! Do not forget to increment `i` otherwise you create an infinite loop! `i` will never be greater than 10 and the condition is always True.. If you are running an infinite loop, you can break it by clicking the stop button next to the `Run` button on the top menu bar. 

### 3.9. `for` loop

The instruction `for` works over a sequence. It is specialized in running over a sequence of several data. We have not yet presented these sequences that we will discuss later however we already know one type of sequence: strings. Strings are sequences. You can therefore run over a string with a `for` loop.

```python
string = "Hi!, my name is Nicolas"
for letter in string:
    print(letter)
```

Here, the variable `letter` takes successively the value of each letter in the string. We then print to the screen each value taken by letter. Since the shell goes to the next line at every `print()` command, one letter is displayed per line. 

### 3.10. Functions II

Functions a used to gather a series of instructions in a block that we can call later on by the name of the function. We have already seen some functions such as `print()` and `type()`. Although there exist a large number of pre-defined functions in Python, it is convenient to write its own (custom) functions. Defining its own function has several advantages such as avoiding to repeat a series of instruction several times in a program. A function is defined as the following:

```python
def name_of_the_function(parameter1, parameter2, parameter3, parameterN):
    # Instruction block
    pass
```

Here is the explanation:
- the keyword `def` which tells Python we start the definition of a function
- the name of the function, same as for a variable. Do not use a name of a variable previuously defined
- the list of the parameters the function will work on, separated by commas and surounded by parenthesis
- the colon symbol to terminate the line
- the instruction block that must be indented. Here we wrote only the special word `pass` that tells Python to read the line and do nothing

The function to write the multiplication table of 7 would be:

```python
def table_of7():
    nb = 7
    i = 0
    while i<10:
        print(i + 1,'*',nb,'=',(i+1)*nb)
        i += 1
```

If we execute this function we get the following.

```python
table_of7()
```

Now let's go a step further and define the function `table` that takes two parameters the corresponding table number it should print and the maximum multiplication factor.

```python
def table(nb, nmax):
    i = 0
    while i < nmax:
        print(i + 1,'*',nb,'=',(i+1)*nb)
        i += 1
```

```python
table(12,14)
```

One can also define default parameters by adding an equal sign followed by the default value next to the parameter in the definition of the function.

```python
def table(nb, nmax=10):
    i = 0
    while i < nmax:
        print(i + 1,'*',nb,'=',(i+1)*nb)
        i += 1
```

```python
table(12)
```

If we get back to the pre-defined functions we have seen before, there is an important difference between let say the function `print()` and `type()`. In non-interactive mode, the later won't return anything (different from writing to the screen) while the former will. Indeed most of the functions work over the parameters, do some computations and return a result. Let's define a `square()` function that returns the square of a number.

```python
def square(nb):
    return nb*nb
```

The instruction `return` tells the function to send back some data that can for example be stored in a variable.

```python
nb = 5
nb2 = square(nb)
print(nb2)
```

Note that you can return more than one variable separated by commas and you can recover these values into several variable also separated by commas.

### 3.11. Lists

In Python lists are objects that can contain others. Those are sequences like `strings` but instead of containing letters, they can contain any object. There are two ways of defining empty lists:

```python
my_list = list()
another_list = []
print(my_list,another_list)
```

You can also create a lists with some elements inside.

```python
my_list = [1,2,3,4]
print(my_list)
```

This list contains 4 object of type `int`. However, lists can contain any object and each object has its own location. To access an element, it is similar than accessing an element in a string by providing the index of the element in the list in brakets. It is also possible to re-assign an element that way.

```python
my_list[1]
```

```python
my_list[1] = "K"
my_list
```

You note that now I have a list of 4 elements of various type, i.e. `int` and `str`.

#### Adding an element at the end of a list
To add an element at the end of a list you can use the function `append`.

```python
my_list.append(12)
my_list
```

We note that functions associated with strings return something whereas those associated with lists do not return anything but they modify the original object (see above `append`).

```python
string1 = "a short sentence"
string2 = string1.upper() # here we uppercase string1
string1
```

```python
string2
```

```python
list1 = [1, 5.5, 18]
list2 = list1.append(-15) 
list1
```

```python
list2
```

Indeed, nothing is displayed from the last cell. If you `print(list2)` you will get `None`, indeed, if a function returns nothing it actually returns `None`. Try it out.

#### Insert an element

You can also insert an element in a list at a given location.

```python
my_list = ['a','b','d','e']
my_list.insert(2,'c')
my_list
```

Above, I tell Python to add element `c` in the list at location index 2. The command will increment indexes greater or equal than 2 and insert `c` at index 2. 

#### Concatenation

You can also increase lists by concatenation of two (or more) lists.

```python
list1 = [1,2,3]
list2 = [4,5,6,7]
list1.extend(list2)
list1
```

```python
list1 = [1,2,3]
list2 = [4,5,6,7]
list1+list2
```

```python
list1 = [1,2,3]
list2 = [4,5,6,7]
list1 += list2
print(list1)
```

Right above are 3 equivalent ways to concatenate lists. Note that `+=` and `extend` are equivalent and do not returnn any results while the operator `+` does return the result.

#### Suppress elements
The `del` function can be used to delete an object such as a variable or an element in a list.

```python
my_variable = 4
del my_variable
my_variable
```

We see that once we have deleted the variable, if we call it, we get an error saying that the variable does not exist.

```python
my_list = [1,2,2,3,4]
del my_list[2]
my_list
```

You ca also use the function `remove` to delete the matching element(s).

```python
my_list = [31,32,38,33,38,34]
my_list.remove(38)
my_list
```

Note that the remove function only removes the first occurence of the element.

#### Walk throught a list

You should already be able to display the elemnents of a list one by one but let see that this can be done in multiple ways.

```python
my_list = ['a','b','c','d','e','f','g','h']
i = 0
while i < len(my_list): # the function len() returns the length of a list, here 8
    print(my_list[i])
    i += 1
```

```python
my_list = ['a','b','c','d','e','f','g','h']
for element in my_list:
    print(element)
```

As you can see, it is similar than what we learnt before about walking through strings. Now let see another way to walk through a list.

```python
my_list = ['a','b','c','d','e','f','g','h']
for i, element in enumerate(my_list):
    print("At the index {} is the element {}.".format(i, element))
```

Here we have a `for` loop but quite special. Between `for` and `in` we have two variables, separated by a comma. In fact, `enumerate` returns an object that is associated with a two variables list: index and element. Let's look at it a bit differently.

```python
my_list = ['a','b','c','d','e','f','g','h']
for elements in enumerate(my_list):
    print(elements)
```

When we look at each elements returned by `enumerate`, we see what's called a `tuple` with 2 elements: the index and the value of the list. Tuples are similar to lists but cannot be modified once created.

#### The function `split()`

This function is very useful when you read some text from a file and need to access some elements in the strings. For example, you often need to remove part of the string to isolate the data you need. The `split()` method splits a string into a list.

```python
string = 'Welcome to AP620'
s = string.split()
print(s)
```

We see that by default (with no argument) the function has separated the string at the space characters. You can extend the feature of this function by providing any splitting character as argument.

```python
string = '163kg'
weight = string.split('kg')
print(weight[0])
```

## 4. Quiz (10 marks)
Answer the following questions by typing the necessary code in the (code) cell below each question instead of `# YOUR CODE HERE`. Ignore the empty cell below the code answer cell, it is used by the autograder.

__Q.1.__ Assign the value 30 to the variable `x` (1 mark).

<!-- grade_id: cell-60f1dcd6ed074240, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
x = 30
### END SOLUTION
```

<!-- grade_id: cell-f8a16e6982cd3c05, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert x == 30
### END HIDDEN TESTS
```

__Q.2.__ Define the 4-element list `squares2` starting with the integer 2 and with each subsequent elements at index `i` ($i>0$) to be the square of the element at index `i-1`. Here no need to write a loop, just define directly the list with the actual numbers (1 mark).

<!-- grade_id: cell-476e4dfad7efee83, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
squares2 = [2,2**2,2**4,2**8]
### END SOLUTION
squares2[3]
```

<!-- grade_id: cell-721511c5465a1cec, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert sum(squares2) == 278
assert squares2[3] == 256
### END HIDDEN TESTS
```

__Q.3.__ Complete the function `within_5_10()` that takes a number as an argument and check if this number is stricly less than 10 and greater or equal than 5. The function must return a boolean i.e. `True` or `False` (2 marks).

<!-- grade_id: cell-285f033d12e8b1db, points: 0, grade: false, solution: true -->
```python
def within_5_10(num):
    ans = False # Here we initialize the boolean to False
    ### BEGIN SOLUTION
    if num >= 5 and num < 10:
        ans = True
    ### END SOLUTION
    return ans
```

<!-- grade_id: cell-e23ea4c1b4ca9310, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert within_5_10(8) == True
assert within_5_10(12) == False
assert within_5_10(5) == True
assert within_5_10(10) == False
### END HIDDEN TESTS
```

__Q.4.__ Implement the function `all_positive()` that takes a list of numbers as an argument and returns `True` if all elements of the list are positive or null, `False` otherwise. You can use a loop `for` to go over the elements of `nums` and check with a `if` statment if it is positive or null. If not, assign `False` to the variable `ans` (2 marks).

<!-- grade_id: cell-a8b5b38743f93260, points: 0, grade: false, solution: true -->
```python
def all_positive(nums):
    ans = True # We initialize the boolean
    ### BEGIN SOLUTION
    for num in nums:
        if num < 0:
            ans = False
    ### END SOLUTION
    return ans
```

<!-- grade_id: cell-909eac762e386505, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert all_positive([2,3,4,5,150]) == True
assert all_positive([1,3,4,5,5.5,-1]) == False
assert all_positive([0,1,2,3,4,5,10.9]) == True
### END HIDDEN TESTS
```

__Q.5.__ Define the function `get_int_index()` that takes a list as input and returns a list of the index of the integer elements. You can also use a for loop however you must track both indices and elements. This can be achieved in various ways. For example you can use a counter to keep track of indices (2 marks).

<!-- grade_id: cell-1794d7241a891679, points: 0, grade: false, solution: true -->
```python
def get_int_index(my_list):
    index_of_int = [] # We initialize and empty list that will contain the indices of integer elements
    ### BEGIN SOLUTION
    for i, element in enumerate(my_list):
        if type(element) == int:
            index_of_int.append(i)
    ### END SOLUTION
    return index_of_int
```

<!-- grade_id: cell-d13158c34fe62c87, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert get_int_index(['a',11,'c','d',14]) == [1,4]
assert get_int_index([10,11,12,'d',14,'15']) == [0,1,2,4]
assert get_int_index(['d','15']) == []
### END HIDDEN TESTS
```

__Q.6.__ Define the function `volume_sphere()` that takes the radius of a sphere in meters as argument and returns its volume in liters (2 marks).

<!-- grade_id: cell-4e2e0502063940cd, points: 0, grade: false, solution: true -->
```python
def volume_sphere(radius):
    pi = 3.14 # use the approximate value of pi
    volume = 0 # initialize the volume
    ### BEGIN SOLUTION
    volume = (4.0/3.0)*pi*radius*radius*radius*1e+3
    ### END SOLUTION
    return volume
```

<!-- grade_id: cell-610465fd112cae85, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert round(volume_sphere(1.0),2) == 4186.67
assert round(volume_sphere(2.0),2) == 33493.33
### END HIDDEN TESTS
```

Don't forget to submit your work before the deadline!!!

Assignment --> lab1 --> Submit